Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ''
COLLABORATORS = ''
```

---

# Домашнее задание №4. Матричная факторизация

## Задачи - 25 баллов (+5 доп баллов)
1. Ноутбук `mf.ipynb` - 20 баллов
- ImplicitALS - 4 балла
- SVD - 4 балла
- Dataset with features - 2 балла
- ImplicitALS with features - 5 баллов
- LightFM with features - 5 баллов
2. Имплементация модели в сервис - 5 баллов
- Пробить на Leaderboard порог `map@10 = 0.075`
- Если при этом используете MF (Implicit или LightFM) + ANN (nmslib, faiss, annoy и тд) - дополнительно 5 баллов
  
## Как сдать ноутбук `mf.ipynb` на проверку

1. Прогоните весь код ноутбука - проверьте, что нет ошибок и тесты проходят
2. Выложите готовый ноутбук в ваш репозиторий с сервисом из домашнего задания №1 по пути `notebooks/hw_4/mf.ipynb` в ветке `hw_4`
3. Проверьте, что есть доступ к вашему репозиторию для аккаунтов `https://github.com/feldlime`
4. Откройте PR в main ветку и добавьте в ревьюеры **своего ментора**
5. Не проводите мердж в `main` ветку, пока не увидите оценку за это ДЗ в ведомости. Файл с ноутбуком должен находиться в ветке `hw_4`

Обратите внимание, что сборка ноутбуков на проверку автоматизирована. В случае неправильного пути, имени файла или ветки (а также при отсутствии доступа у `@feldlime`) ваша работа не попадёт на проверку и получит `0` баллов.

Используемые библиотеки в рамках ДЗ
```bash
pip install implicit==0.7.2 requests==2.32.3 rectools[lightfm]==0.12.0 pandas==2.2.3 numpy==1.26.4 scipy==1.12.0
```

## Импорты и данные

```python
import os
import os.path
import pickle

import numpy as np
import pandas as pd
import threadpoolctl
from rectools import Columns
from rectools.dataset import Dataset
from rectools.metrics import MAP
from rectools.models import model_from_config
from rectools.tools import UserToItemAnnRecommender

# For implicit ALS
os.environ['OPENBLAS_NUM_THREADS'] = '1'
threadpoolctl.threadpool_limits(1, 'blas')
```

Output:
```
d:\GitHub\itmo_recsys\.venv\Lib\site-packages\tqdm\auto.py:21: TqdmWarning: IProgress not found. Please update jupyter and ipywidgets. See https://ipywidgets.readthedocs.io/en/stable/user_install.html
  from .autonotebook import tqdm as notebook_tqdm
d:\GitHub\itmo_recsys\.venv\Lib\site-packages\lightfm\_lightfm_fast.py:9: UserWarning: LightFM was compiled without OpenMP support. Only a single thread will be used.
  warnings.warn(
<threadpoolctl.threadpool_limits at 0x2bffe710410>
```

Если вдруг у вас нет данных, то используйте закомментированный код

```python
# url = 'https://github.com/irsafilo/KION_DATASET/raw/f69775be31fa5779907cf0a92ddedb70037fb5ae/data_original.zip'

# req = requests.get(url, stream=True)

# with open('kion.zip', 'wb') as fd:
#     total_size_in_bytes = int(req.headers.get('Content-Length', 0))
#     progress_bar = tqdm(desc='kion dataset download', total=total_size_in_bytes, unit='iB', unit_scale=True)
#     for chunk in req.iter_content(chunk_size=2 ** 20):
#         progress_bar.update(len(chunk))
#         fd.write(chunk)

# files = zf.ZipFile('kion.zip', 'r')
# files.extractall()
# files.close()
```

```python
data_path = os.environ.get('DATA_PATH')


if data_path is None:
    data_path = '../data_original'


interactions = pd.read_csv(
    os.path.join(data_path, 'interactions.csv'), parse_dates=['last_watch_dt']
)
interactions = interactions.rename(
    columns={'total_dur': Columns.Weight, 'last_watch_dt': Columns.Datetime}
)

users = pd.read_csv(os.path.join(data_path, 'users.csv'))
items = pd.read_csv(os.path.join(data_path, 'items.csv'))

print(interactions.shape)
interactions.head()
```

Output:
```
(5476251, 5)
   user_id  item_id   datetime  weight  watched_pct
0   176549     9506 2021-05-11    4250         72.0
1   699317     1659 2021-05-29    8317        100.0
2   656683     7107 2021-05-09      10          0.0
3   864613     7638 2021-07-05   14483        100.0
4   964868     9506 2021-04-30    6725        100.0
```

```python
N_DAYS = 7

max_date = interactions['datetime'].max()
train = interactions[(interactions['datetime'] <= max_date - pd.Timedelta(days=N_DAYS))]
test = interactions[(interactions['datetime'] > max_date - pd.Timedelta(days=N_DAYS))]

catalog = train[Columns.Item].unique()

test_users = test[Columns.User].unique()
cold_users = set(test_users) - set(train[Columns.User])
test.drop(test[test[Columns.User].isin(cold_users)].index, inplace=True)
hot_users = test[Columns.User].unique()

train.shape, test.shape, len(catalog), len(test_users), len(hot_users), len(cold_users)
```

Output:
```
C:\Users\iqpiv\AppData\Local\Temp\ipykernel_20356\3411243360.py:11: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  test.drop(test[test[Columns.User].isin(cold_users)].index, inplace=True)
((5051815, 5), (306752, 5), 15577, 167348, 111240, 56108)
```

```python
K_RECOS = 10
map10 = MAP(k=K_RECOS)
```

## ImplicitALS

### Ситуация:

Коллега вернулся из отпуска и вы вместе сели за улучшение модели. Внимательно изучив репозиторий библиотеки implicit вы увидели модель iALS и решаете попробовать ее в деле.

Чтобы работа была интереснее, вы заключаете пари с вашим коллегой о том, кто выбьет больше MAP@K на горячих пользователях. 

Правила пари: 
- Валидируемся на последней неделе (переменная `test`) и на горячих пользователях `hot_users`
- Можно собрать свой `Dataset` на основе `train`, трансформированного, если нужно
- Параметры модели задаются конфигом, которые будут передаваться в `model_from_config`

У вашего коллеги получилось выбить на ImplicitALS `MAP@K = 0.052`. Ваша задача побить его рекорд.

```python
train_2 = train.copy()
train_2[Columns.Weight] = np.where(train_2[Columns.Weight] > 10, 10, 1)
dataset = Dataset.construct(train_2)
```

```python
config = {
    'cls': 'ImplicitALSWrapperModel',
    'model': {
        'factors': 4,
        'regularization': 0.5,
        'alpha': 5,
        'iterations': 64,
        'use_gpu': False,
        'random_state': 42,
    },
}
```

<!-- grade_id: cell-6c041ad93abfab9a, points: 4, grade: true, solution: false -->
```python
# %%time
assert config['cls'] == 'ImplicitALSWrapperModel'

model = model_from_config(config)
model.fit(dataset)

recos = model.recommend(
    users=hot_users,
    dataset=dataset,
    k=K_RECOS,
    filter_viewed=True,
)
print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.052
```

Output:
```
0.06828654560913162
```

## SVD

На ваш громкий спор с коллегой о том, что все дело в вашем удачном random seed, к вам подошел ваш лид. 

Узнав детали вашего спора, он дает вам комментарий, что iALS хороша, но погружение в матричную факторизацию следует начинать с `SVD`.

Вы переглянулись с коллегой и решаете уладить спор о random seed во втором раунде, используя новую модель.

Ваш коллега смог выбить на SVD `MAP@K = 0.066`. Вы знаете, что делать.

<!-- grade_id: cell-f1b963c73db09e3a, points: 0, grade: false, solution: true -->
```python
config = {
    'cls': 'PureSVDModel',
    'factors': 2,
    'use_gpu': False,
    'random_state': 42,
}
```

<!-- grade_id: cell-366e2b318e3b502d, points: 4, grade: true, solution: false -->
```python
# %%time
assert config['cls'] == 'PureSVDModel'

model = model_from_config(config)
model.fit(dataset)

recos = model.recommend(
    users=hot_users,
    dataset=dataset,
    k=K_RECOS,
    filter_viewed=True,
)
print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.066
```

Output:
```
0.0703870578041766
```

## Dataset with features

"Ну это ни в какие ворота!" - восклицает ваш коллега, увидев ваш победный конфиг. Из другого угла опенспейса доносится "А я говорил" от вашего лида.

В это время к вам сзади подходит продакт и интересуется предметом вашего спора.

Рассказав про особенности найденных вами моделей, он просит вас в них докинуть фичи, ведь на одних взаимодействиях далеко не уедешь.

Вы согласились, ведь это отличная возможность продолжить пари. Соберите `Dataset` с фичами по пользователям и итемам, который вы будете использовать дальше


```python
users = users.loc[users[Columns.User].isin(interactions[Columns.User])].copy()
user_features_frames = []

for feature in ['sex', 'age', 'income']:
    feature_frame = users.reindex(columns=[Columns.User, feature])
    feature_frame.columns = ['id', 'value']
    feature_frame['feature'] = feature
    user_features_frames.append(feature_frame)

user_features = pd.concat(user_features_frames)
user_features.tail()
```

Output:
```
id         value feature
840189  191349  income_40_60  income
840190  393868  income_20_40  income
840192  339025   income_0_20  income
840194  251008           NaN  income
840195  590706           NaN  income
```

```python
items = items.loc[items[Columns.Item].isin(interactions[Columns.Item])].copy()

items['genre'] = (
    items['genres'].str.lower().str.replace(', ', ',', regex=False).str.split(',')
)
genre_feature = items[['item_id', 'genre']].explode('genre')
genre_feature.columns = ['id', 'value']
genre_feature['feature'] = 'genre'

content_feature = items.reindex(columns=[Columns.Item, 'content_type'])
content_feature.columns = ['id', 'value']
content_feature['feature'] = 'content_type'

item_features = pd.concat((genre_feature, content_feature))
item_features.tail()
```

Output:
```
id   value       feature
15958   6443  series  content_type
15959   2367  series  content_type
15960  10632  series  content_type
15961   4538  series  content_type
15962   3206  series  content_type
```

<!-- grade_id: cell-b0397a70b8235dd9, points: 0, grade: false, solution: true -->
```python
dataset_with_features = Dataset.construct(
    interactions_df=train_2,
    user_features_df=user_features,
    cat_user_features=['sex', 'age', 'income'],
    item_features_df=item_features,
    cat_item_features=['genre', 'content_type'],
)
```

<!-- grade_id: cell-7773b657b67af87b, points: 2, grade: true, solution: false -->
```python
assert (dataset_with_features.user_features is not None) and (
    dataset_with_features.item_features is not None
)
```

## ImplicitALS with features

Собрав датасет с фичами вы готовы к третьему раунду пари. 

Вы решаете начать снова с `iALS`, до сих пор удивляясь результатам модели `SVD`.

Ваш коллега изучил вашу технику подбора random seed и хитро улыбается вам.

Он смог выбить `MAP@K = 0.073`, теперь ваш ход.

<!-- grade_id: cell-651def1b6c1a52b6, points: 0, grade: false, solution: true -->
```python
config = {
    'cls': 'ImplicitALSWrapperModel',
    'model': {
        'factors': 4,
        'regularization': 0.5,
        'alpha': 5,
        'iterations': 64,
        'use_gpu': False,
        'random_state': 42,
    },
}
```

<!-- grade_id: cell-f3ed67874a568322, points: 5, grade: true, solution: false -->
```python
# %%time
assert config['cls'] == 'ImplicitALSWrapperModel'

model = model_from_config(config)
model.fit(dataset_with_features)

recos = model.recommend(
    users=hot_users,
    dataset=dataset_with_features,
    k=K_RECOS,
    filter_viewed=True,
)
print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.073
```

Output:
```
d:\GitHub\itmo_recsys\.venv\Lib\site-packages\rectools\dataset\features.py:424: UserWarning: Converting sparse features to dense array may cause MemoryError
  warnings.warn("Converting sparse features to dense array may cause MemoryError")
100%|██████████| 1/1 [00:03<00:00,  3.98s/it]
100%|██████████| 1/1 [00:05<00:00,  5.12s/it]
0.08206192770527554
```

## LightFM with features

И снова ор выше гор, ваш пайплайн подготовки датасета помог вам в очередной раз обойти вашего коллегу.

Не зная, к чему еще аппелировать, он зовет вашего старшего коллегу, чтобы тот внимательно изучил полученные результаты.

"iALS с фичами это хорошо, но тут стоит попробовать факторизационные машины, попробуйте `LightFM`" - заключает он. Вы переключаетесь на изучение новой библиотеки, предвкушая финальный раунд.

Ваш коллега смог выжать из своего обновленного `Dataset` и `LightFM` скор `MAP@10 = 0.08`. Последний рывок.

```python
config = {
    'cls': 'LightFMWrapperModel',
    'model': {
        'loss': 'warp',
        'no_components': 100,
        'learning_rate': 0.0004069497441886399,
        'user_alpha': 0.5,
        'item_alpha': 0.88,
        'random_state': 42,
    },
}
```

<!-- grade_id: cell-754f9152c2a343a5, points: 5, grade: true, solution: false -->
```python
# %%time
assert config['cls'] == 'LightFMWrapperModel'

model = model_from_config(config)
model.fit(dataset_with_features)

recos = model.recommend(
    users=hot_users,
    dataset=dataset_with_features,
    k=K_RECOS,
    filter_viewed=True,
)
print(map10.calc(recos, test))

assert map10.calc(recos, test) >= 0.08
```

Output:
```
0.08096089312519526
```

## Сервис

Эта битва была легендарной, но вот за спиной опять появляется ваш продакт. 

"Время катить АБ" - говорит он в пятницу вечером. Делать нечего, вы собираете ваши наработки и определяете совместный фронт работ.

В прод должна заехать лучшая модель, которая побьет текущей модели в проде в `MAP@10 = 0.075`.

Также есть бонус от вашего старшего коллеги, который вам советует присмотреться к `Approximate nearest neighbours`, например к `nmslib`.

"Если сможешь обернуть в ANN, то на следующем годовом ревью получишь от меня оценку отлично" - сказал он. Изучить новые технологии и получить повышение - идеально, заключаете вы и бросаетесь в бой.

# Вместо заключения

## Задачи - 25 баллов (30 баллов с доп задачей по ANN)
1. Ноутбук `mf.ipynb` - 20 баллов
- SVD - 5 баллов
- ImplicitALS - 5 баллов
- ImplicitALS with features - 5 баллов
- LightFM with features - 5 баллов
2. Имплементация модели в сервис - 5 баллов
- Пробить на Leaderboard порог `map@10 = 0.075`
- Если при этом используете MF (Implicit или LightFM) + ANN (nmslib, faiss, annoy и тд) - дополнительно 5 баллов
  
## Как сдать ноутбук `mf.ipynb` на проверку

1. Прогоните весь код ноутбука - проверьте, что нет ошибок и тесты проходят
2. Выложите готовый ноутбук в ваш репозиторий с сервисом из домашнего задания №1 по пути `notebooks/hw_4/mf.ipynb` в ветке `hw_4`
3. Проверьте, что есть доступ к вашему репозиторию для аккаунтов `https://github.com/feldlime`
4. Откройте PR в main ветку и добавьте в ревьюеры **своего ментора**
5. Не проводите мердж в `main` ветку, пока не увидите оценку за это ДЗ в ведомости. Файл с ноутбуком должен находиться в ветке `hw_4`

Обратите внимание, что сборка ноутбуков на проверку автоматизирована. В случае неправильного пути, имени файла или ветки (а также при отсутствии доступа у `@feldlime`) ваша работа не попадёт на проверку и получит `0` баллов.

```python
interactions[Columns.Weight] = np.where(interactions[Columns.Weight] > 10, 10, 1)

dataset_with_features = Dataset.construct(
    interactions_df=interactions,
    user_features_df=user_features,
    cat_user_features=['sex', 'age', 'income'],
    item_features_df=item_features,
    cat_item_features=['genre', 'content_type'],
)

config = {
    'cls': 'ImplicitALSWrapperModel',
    'model': {
        'factors': 4,
        'regularization': 0.5,
        'alpha': 5,
        'iterations': 64,
        'use_gpu': False,
        'random_state': 42,
    },
}

model = model_from_config(config)
model.fit(dataset_with_features)

```

Output:
```
d:\GitHub\itmo_recsys\.venv\Lib\site-packages\rectools\dataset\features.py:424: UserWarning: Converting sparse features to dense array may cause MemoryError
  warnings.warn("Converting sparse features to dense array may cause MemoryError")
100%|██████████| 1/1 [00:02<00:00,  2.65s/it]
100%|██████████| 1/1 [00:05<00:00,  5.49s/it]
<rectools.models.implicit_als.ImplicitALSWrapperModel at 0x2bff210c410>
```

```python
user_vectors, item_vectors = model.get_vectors()

model_ann = UserToItemAnnRecommender(
    user_vectors=user_vectors,
    item_vectors=item_vectors,
    user_id_map=dataset_with_features.user_id_map,
    item_id_map=dataset_with_features.item_id_map,
)
model_ann.fit()
```

Output:
```
<rectools.tools.ann.UserToItemAnnRecommender at 0x2bf95f0e250>
```

```python
with open('../../models/als_with_features_model.pickle', 'wb') as file:
    pickle.dump(model, file)

model_ann.index.saveIndex('../../models/als_ann_with_features_index.pickle')

with open('../../models/als_ann_dataset_user_id_map.pickle', 'wb') as file:
    pickle.dump(dataset_with_features.user_id_map, file)

with open('../../models/als_ann_dataset_item_id_map.pickle', 'wb') as file:
    pickle.dump(dataset_with_features.item_id_map, file)
```